---
title       : Insert the chapter title here
description : Insert the chapter description here


--- type:NormalExercise lang:r xp:100 skills:1 key:53eb2ef260
## Test


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:69ba51cb36
## Test

*** =slides_key
c1184800af6af5beb845a9b88f0d6857
